## Week 5 Assignment - W200 Python for Data Science, UC Berkeley MIDS

Write code in this jupyter Notebook to solve each of the following problems. Each problem should have its solution in a separate cell. Please upload this **Notebook** with your solutions to your GitHub repository in your SUBMISSIONS/week_05 folder by 11:59PM PST the night before class.

## Objectives:

- Demonstrate how to define functions 
- Understand how to call functions from within other functions
- Understand how to return results from a function
- Use the Python namespace to configure variables
- Use functions as objects to pass to other functions
- Design and implement a recursive function
- Demonstrate knowledge of the Try/Except error checking statements

## General Guidelines:

- All calculations need to be done in the functions (that includes any formatting)
- Name your functions exactly as written in the problem statement
- Please have your functions return the answer rather than printing it inside the function
- Do not make a separate input() statement.  The functions will be passed the input as shown in the examples
- The examples given are samples of how we will test/grade your code. Please ensure your functions output the same information
- There are also test cells for you to test your functions against the examples and the autograder.
- Answer format is graded - please match the examples
- Other than for Part 5, user / function inputs do not need to be validated or checked. (For example, if the problem states input an integer we will check it by inputting an integer)
- Docstrings and comments in your code are strongly suggested but won't be graded
- This homework is mostly auto-graded.  The blank code blocks are the auto-grading scripts - please do not delete these!
- Your code needs to be written in the #Your Code Here blocks or it wont be graded correctly.

## 5-1. Function Exercises

### 5-1-1. Nested ("Wrapped") Functions
For this question, please write three functions as follows:

- `sum_digits`:   a function which takes an `int` and returns the sum of its (positive value) digits. 

- `diff_sum_digits`:   a function that "wraps" `sum_digits` in that it calls `sum_digits` from within it. Use the `diff_sum_digits` function to compute the absolute value of the input number minus the sum of digits of the input number.

- `wraps_diff_sum_digits`:   a function that "wraps" `diff_sum_digits`. If `diff_sum_digits` returns a result that has more than one digit have this new function replace the result with the sum of the digits of the result. Do this repeatedly until the result has just one digit, then display it.

To illustrate this with an example:
- The input number is 20 as in: wraps_diff_sum_digits(20)
- wraps_diff_sum_digits calls diff_sum_digits(20) which calls sum_digits(20)
- sum_digits adds the numbers (2+0 = 2) returns a 2, diff_sum_digits subtracts (20-2 = 18) and returns this to wraps_diff_sum_digits
- wraps_diff_sum_digits sees that the number 18 stills has 2 digits and calls diff_sum_digits(18) which calls sum_digits(18)
- sum_digits add the numbers (1+8 = 9) returns a 9, diff_sum_digits subtracts (18-9 = 9) and returns this to wraps_diff_sum_digits
- wraps_diff_sum_digits sees that the number 9 only has 1 digit - stops and returns the value 9.

Below you'll find an example of what we mean when we say "wraps".

```python
def example_base_func(x):
    "Returns the value of the input * -1"
    return -1 * x

def wraps_example_base_func(x):
    temp_val = example_base_func(x)
    if temp_val < 0:
        return "trivial example"
    
print(example_base_func(5))
print(wraps_example_base_func(5))
```

<!-- grade_id: cell-c373bf43a4b7ada6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def sum_digits(x):
    """a function which takes an int and returns the sum of its (positive value) digits."""
    x = abs(x)
    result = 0
    while x >= 1:
        result += x%10
        x //= 10
    
    return result

def diff_sum_digits(x):
    """a function that "wraps" sum_digits in that it calls sum_digits from within it. 
    Use the diff_sum_digits function to compute the absolute value of the input number minus the sum of digits of the input number."""
    return abs(x - sum_digits(x))

def wraps_diff_sum_digits(x):
    result = diff_sum_digits(x)
    while result >= 10:
        result = diff_sum_digits(result)
        
    return result

number = int(input("Please enter an integer: "))

print(wraps_diff_sum_digits(number))
        
```

Output:
```
Please enter an integer: 54321
15
54306
9
```

You can use the sample output in below cell to sanity check your work.

We also show you some of the tests we will run below.

```python
# You can use the sample output in this cell to sanity check your work.

print(sum_digits(54321) == 15)
print(sum_digits(-54321) == 15)
print(diff_sum_digits(54321) == 54306)
print(wraps_diff_sum_digits(54321) == 9)
```

Output:
```
True
True
True
True
```

<!-- grade_id: cell-88242cbaef2afe19, points: 1, grade: true, solution: false -->
```python
# Autograde cell - do not erase/delete
# Below are some of the autograding tests we run to give you an idea on how they work.
from nose.tools import assert_equal
assert_equal(sum_digits(54321), 15)
```

<!-- grade_id: cell-69021db4737eb29b, points: 1, grade: true, solution: false -->
```python
# Autograde cell - do not erase/delete
assert_equal(sum_digits(-54321), 15)
```

<!-- grade_id: cell-257a5430df922fd1, points: 1, grade: true, solution: false -->
```python
# Autograde cell - do not erase/delete
assert_equal(diff_sum_digits(54321), 54306)
```

<!-- grade_id: cell-90fd0a6f7ef73ec9, points: 2, grade: true, solution: false -->
```python
# Autograde cell - do not erase/delete
assert_equal(wraps_diff_sum_digits(54321), 9)
```

<!-- grade_id: cell-1a7c387b6af6d62d, points: 2, grade: true, solution: false -->
```python
# Autograde cell - do not erase/delete
# These are the hidden autograder cells that will be run to check your functions!
```

<!-- grade_id: cell-58c35e06fc9646e9, points: 2, grade: true, solution: false -->
```python
# Autograde cell - do not erase/delete
```

<!-- grade_id: cell-1f13c56183fc697f, points: 1, grade: true, solution: false -->
```python
# Autograde cell - do not erase/delete
```

### 5-1-2. Pigs, Continued

Write a function `is_consonant` that takes a character and returns `True` if it is a consonant.

Use your function to create a new function `to_piglatin` that takes a word, moves all starting consonants (all consonants before the first vowel) to the end of the word, then adds *ay* to the end and returns the result. You may expect that the input to the function will be just one word. (we know this isn't **true** pig latin - please do not change this basic algorithm). For a single word input the first letter is capitalized and the rest are lower case as shown in the example below.

*Extra Credit (5 points): Have the 'to_piglatin' function check whether or not the input is multiple words and return the whole sentence in pig latin. The sentence will be returned in sentence case (only the first word capitalized). For this you may assume that a sentence always ends with a period and the input is always one sentence, never more than one. Make these changes in the 'to_piglatin' function - that is, if you want to do the extra credit the 'to_piglatin' function needs to work for both a single word and a sentence. NOT AS A SEPARATE FUNCTION*

Examples:
```
(format is: function call -> returns) - do not return this whole string just return the word after the arrow.

to_piglatin('stay') ->  Aystay
to_piglatin('Jared') -> Aredjay 
to_piglatin('and') -> Anday
to_piglatin('CAR') -> Arcay
```

<!-- grade_id: cell-2eef38fa43df4ac5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def is_consonant(char):
    if char.lower() not in "aeiou":
        return True
    
def to_piglatin(words):
    words_string = words.lower().replace(".","").split()
    for k in range(0,len(words_string)): 
        l = len(words_string[k])
        while is_consonant(words_string[k][0]):
            words_string[k] = words_string[k][1:l] + words_string[k][0]
        words_string[k] = words_string[k] + "ay"
        if k == 0:
            words_string[k] = words_string[k].capitalize()
    
    if len(words_string) > 1:
        result = " ".join(words_string) + "."
    else:
        result = " ".join(words_string)
    
    return result
            
words = input("Please enter a word or a sentence: ")
print(to_piglatin(words))

```

Output:
```
Please enter a word or a sentence: Stay away from me.
Aystay awayay omfray emay.
```

```python
# You can use the sample output in this cell to sanity check your work.

print(to_piglatin("stay") == "Aystay")
print(to_piglatin("Jared") == "Aredjay")
print(to_piglatin("and") == "Anday")
print(to_piglatin("CAR") == "Arcay")
```

Output:
```
True
True
True
True
```

<!-- grade_id: cell-8f844f982b02b98c, points: 0, grade: true, solution: false -->
```python
# Autograde cell - do not erase/delete
assert_equal(to_piglatin("stay"), "Aystay")
```

<!-- grade_id: cell-aa09ec12198816d7, points: 0, grade: true, solution: false -->
```python
# Autograde cell - do not erase/delete
assert_equal(to_piglatin("Jared"), "Aredjay")
```

<!-- grade_id: cell-85e4811f3f7ac483, points: 0, grade: true, solution: false -->
```python
# Autograde cell - do not erase/delete
assert_equal(to_piglatin("and"), "Anday")
```

<!-- grade_id: cell-f6d83d317209a5c5, points: 0, grade: true, solution: false -->
```python
# Autograde cell - do not erase/delete
assert_equal(to_piglatin("CAR"), "Arcay")
```

<!-- grade_id: cell-5ee412cb4b703023, points: 2, grade: true, solution: false -->
```python
# Hidden autograde cell - do not erase/delete
```

<!-- grade_id: cell-954080f6af7c07aa, points: 2, grade: true, solution: false -->
```python
# Hidden autograde cell - do not erase/delete
```

<!-- grade_id: cell-68acd2e1a791ca95, points: 2, grade: true, solution: false -->
```python
# Hidden autograde cell - do not erase/delete
```

<!-- grade_id: cell-e39d29342e27b90c, points: 2, grade: true, solution: false -->
```python
# Hidden autograde cell - do not erase/delete
```

<!-- grade_id: cell-5512005b0518cb19, points: 2, grade: true, solution: false -->
```python
# Hidden autograde cell - do not erase/delete
```

<!-- grade_id: cell-ba4425be60b298fb, points: 5, grade: true, solution: false -->
```python
# Hidden autograde cell - do not erase/delete
```

## 5-2. Namespace Exercise

### 5-2-1. Process Order

The following code defines a customer order list.   Each item in the list is a tuple containing the name of the item, the number requested, and the price per each.  

Fill in just the `process_order` function below without changing any other code. This function should iterate through all items in the list, removing one item at a time and printing a nice message about it (from within the function).  Finally, confirm that the final print statement correctly displays the total price for all the items.

Example:
```
You filled an order for 1 antacid for a total of $5.33
You filled an order for 3 sour bites for a total of $6.99
You filled an order for 1 gummy bears for a total of $1.99
You filled an order for 4 oranges for a total of $12.88
Total price: $27.19
```

<!-- grade_id: cell-74a84f91e3b3e716, points: 0, grade: false, solution: true -->
```python
### [Youzhi Wu] 

total = 0

def process_order(x_list):
    # YOUR CODE HERE
    global total
    while len(x_list)>0:
        item, num, unit_price = x_list[0][0], x_list[0][1], x_list[0][2]
        item_total = num * unit_price
        print("You filled an order for", num, item, "for a total of", "${:.2f}".format(item_total))
        total += item_total
        x_list.pop(0)
    

######################################################################################################
# DO NOT CHANGE ANY OF THE CODE BELOW HERE                                                           #
######################################################################################################

x = [("oranges", 4, 3.22),("gummy bears",1,1.99),("sour bites", 3, 2.33), ("antacid", 1, 5.33)]

while(len(x)>0):
    process_order(x)

print("Total price: ${:.2f}".format(total))
```

Output:
```
You filled an order for 4 oranges for a total of $12.88
You filled an order for 1 gummy bears for a total of $1.99
You filled an order for 3 sour bites for a total of $6.99
You filled an order for 1 antacid for a total of $5.33
Total price: $27.19
```

```python
# You can use the sample output in this cell to sanity check the end result in your work.
# However, check also your intermediary steps!

total = 0

x = [("oranges", 4, 3.22),("gummy bears",1,1.99),("sour bites", 3, 2.33), ("antacid", 1, 5.33)]

while(len(x)>0):
    process_order(x)

print("Total price: ${:.2f}".format(total) == 'Total price: $27.19')

```

Output:
```
You filled an order for 4 oranges for a total of $12.88
You filled an order for 1 gummy bears for a total of $1.99
You filled an order for 3 sour bites for a total of $6.99
You filled an order for 1 antacid for a total of $5.33
True
```

<!-- grade_id: cell-b7c581157c7c15f0, points: 20, grade: true, solution: false -->
```python
# Hidden autograde cell - do not erase/delete
```

## 5-3. Functions as Objects

### 5-3-1. A Flexible "Scoring" Function

The following code defines a list of names and also contains a header for the function `best`.  The `best` function takes two arguments: a generic scoring function, score, and a list of strings, names.  Fill in the function so that it applies a score function to each string in the names list and returns the name with the highest score. If there are ties in the list, your function should return the first item with the maximum score. The `best` function needs to be designed so that it can take any scoring function and return the name with the highest score.

For this question, define a function called `len_score` that returns the length of a word.  Call the `best` function with the `len_score` function as a parameter.

Example:
```
print(best(len_score, names), "has the longest name.")
    
McJagger has the longest name.
```

<!-- grade_id: cell-08cb8d30cb563aba, points: 0, grade: false, solution: true -->
```python
### [Youzhi Wu]
# YOUR CODE HERE
def len_score(word):
    return len(word)

def best(score,names):
    max_result = 0
    for item in names:
        score_result = score(item)
        if max_result < score_result:
            max_result = score_result
            max_item = item
        
    return max_item

        
```

```python
# You can use the sample output in this cell to sanity check your work.

names = ["Ben", "April", "Zaber", "Alexis", "McJagger", "J.J.", "Madonna"]

print(best(len_score, names) == 'McJagger')

```

Output:
```
True
```

<!-- grade_id: cell-71f0ea37a7c2e641, points: 0, grade: true, solution: false -->
```python
# Autograde cell - do not erase/delete
# Below are some tests we run.

names = ["Ben", "April", "Zaber", "Alexis", "McJagger", "J.J.", "Madonna"]

assert_equal(best(len_score, names), 'McJagger')
```

<!-- grade_id: cell-31e3eba0fe34124e, points: 5, grade: true, solution: false -->
```python
# Hidden autograde cell - do not erase/delete!
```

### 5-3-2. Using Our Flexible Function

**NOTE: Do not change your best function from the answer in 5-3-1 above! The best function should be able to take any scoring function and return the result.** (you also do not need to copy your best function from above - we will run the cell above and then this cell in succession) 

Define a function, `number_of_vowels`, that returns the number of vowels in a string.  Use it with your `best` function to find the name in `names` with the most vowels.

See how easy it is to change the score function for different functionalities!

Example: 
```
print(best(number_of_vowels, names), "has the most vowels.")

Alexis has the most vowels.
```

<!-- grade_id: cell-be845cc598786484, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def number_of_vowels(word):
    count = 0
    for char in word:
        if char.lower() in "aeiou":
            count += 1
    return count
```

```python
# You can use the sample output in this cell to sanity check your work.
names = ["Ben", "April", "Zaber", "Alexis", "McJagger", "J.J.", "Madonna"]

print(best(number_of_vowels, names) == 'Alexis')
```

Output:
```
True
```

<!-- grade_id: cell-6a81391922840599, points: 0, grade: true, solution: false -->
```python
# Autograde cell - do not erase/delete
# Below are some tests we run.

names = ["Ben", "April", "Zaber", "Alexis", "McJagger", "J.J.", "Madonna"]

assert_equal(best(number_of_vowels, names), 'Alexis')
```

<!-- grade_id: cell-5223c7b3911dd60c, points: 5, grade: true, solution: false -->
```python
# Hidden autograde cell - do not erase/delete
```

### 5-3-3 Using Our Flexible Function with a Lambda Function 

**NOTE: Do not change your best function from the answer in 5-3-1 above! The best function should be able to take any scoring function and return the result.** (you also do not need to copy your best function from above - we will run the cells above and then this cell in succession) 

Now pass a `lambda` function into your `best` function to find the name in `names` with the highest number of A's (upper or lower case).  This needs to be a lambda function and should only be one line!

Example output: 
```
Madonna has the most a's.
```

<!-- grade_id: cell-8b3be5cd496fcbe5, points: 10, grade: true, solution: false -->
```python
# Hidden autograde cell - do not erase/delete
```

<!-- grade_id: cell-8e077c8a6f05a83c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
names = ["Ben", "April", "Zaber", "Alexis", "McJagger", "J.J.", "Madonna"]

print(best(lambda x: x.lower().count("a"), names), "has the most a's.")

```

Output:
```
Madonna has the most a's.
```

We cannot help you in sanity checking you work, as you must pass the `lambda` function directly into your `best` function. But you can easily sanity check your solution yourself. 

## 5-4. Recursion Basics

### 5-4-1. Computing Fibonacci Numbers

You are probably familiar (from two homework ago!) with the famous Fibonacci sequence of numbers, which begins like this:
```
1, 1, 2, 3, 5, 8, 13, 21...
```
We'll index from 0, so the 0th and 1st numbers are both 1.  The 2nd Fibonacci number is found by summing the 0th and 1st: 1 + 1 = 2.  The 3rd is found by summing the 1st and 2nd: 1 + 2 = 3.  After this point, each Fibonacci number is found by summing the previous two numbers.

You are to write a recursive function to compute the nth Fibonacci number.  This means that your function will call itself and will **NOT** include explicit loops.

Hint: Your function should include a line that looks a lot like the mathematical definition of the nth Fibonacci number.

Another Hint: It's possible for a recursive function to call itself more than once.

Example: 
```
print(Fibonacci(1))

1
```

<!-- grade_id: cell-908240fb166237b7, points: 0, grade: false, solution: true -->
```python
def Fibonacci(n):
    # YOUR CODE HERE
    if n == 1 or n == 2:
        return 1
    else:
        return Fibonacci(n-2) + Fibonacci(n-1)

#print(Fibonacci(2), Fibonacci(3), Fibonacci(5), Fibonacci(7))
```

Output:
```
1 2 5 13
```

```python
# You can use the sample output in this cell to sanity check your work.

print(Fibonacci(1) == 1)
```

Output:
```
True
```

<!-- grade_id: cell-f99e3713f50786fc, points: 0, grade: true, solution: false -->
```python
# Autograde cell - do not erase/delete
# Below are some tests we run.

assert_equal(Fibonacci(1), 1)
```

<!-- grade_id: cell-dc54b9bfecb04cf3, points: 5, grade: true, solution: false -->
```python
# Hidden autograde cell - do not erase/delete
```

<!-- grade_id: cell-51d06362f92a61dd, points: 5, grade: true, solution: false -->
```python
# Hidden autograde cell - do not erase/delete
```

<!-- grade_id: cell-751b7f716567b5ce, points: 5, grade: true, solution: false -->
```python
# Hidden autograde cell - do not erase/delete
```

### 5-4-2. Computing Fibonacci Numbers (part 2)

Once you are done, think about whether your recursive function is algorithmically efficient.  Compare it to a non-recursive implementation (from homework 3). Which algorithm is more efficient? Why? Please fill in your answer below.

<!-- grade_id: cell-aa13d77f4cd1437c, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

In fact the recursive function is less algorithmically efficient than the non-recursive loop implementation. Take Fibonacci(6) as an example to illustrate:
1. In the recursive function, to return Fibonacci(6), the function will first compute Fibonacci(5) and then Fibonacci(4). For Fibonacci(5) = Fibonacci(3) + Fibonacci(4); Fibonacci(4) = Fibonacci(3) + Fibonacci(2), etc. You can see that in order to get Fibonacci(6), Fibonacci(4), Fibonacci(3), Fibonacci(2), and Fibonacci(1) are computed multiple times. 
2. In the non-recursive function, they would only be computed once. 

## 5-5. Exceptions Exercises

### 5-5-1. Re-factoring for Exceptions

Change the code below that attempts to compute a reciprocal.  It should keep prompting the user for a number until a valid number is entered.  If the user enters 0, the program should print that zero does not have a reciprocal and then terminate.

<!-- grade_id: cell-c78e5f25c8aff22e, points: 10, grade: true, solution: true -->
```python
while True:
    try:
        x = float(input("Enter a number: "))
        print("The reciprocal of your number is", 1/x)
        break
    except ValueError:
        print("You did not enter a valid number.")
        continue
    except ZeroDivisionError:
        print("Zero does not have a reciprocal")
        break
    except:
        print("something else went wrong.")
        continue

```

Output:
```
Enter a number: wi
You did not enter a valid number.
Enter a number: o
You did not enter a valid number.
Enter a number: -owi
You did not enter a valid number.
Enter a number: 8
The reciprocal of your number is 0.125
```

We cannot help you sanity check your code, but you can easily check your solution yourself. What non-valid entries could there be and how should your code behave?

<!-- grade_id: cell-92457557ef030651, points: 0, grade: true, solution: false -->
```python
# Hidden autograde cell - do not erase/delete
```

### 5-5-2 Raising Custom Exceptions

Write a function called `weighted_avg` that takes a list of grades and a corresponding list of weights and returns the weighted average of the grades.  Your function should raise an exception if:
- a weight is less than 0 or greater than 100
- the weights do not add to 100
- the number of weights and grades are not equal
- a grade is below 0 (grades above 100 would be considered extra credit and are acceptable)

Run your function on `grades1` with `weights1` and `grades2` with `weights2` and `grades3` with `weights3` and `grades4` with `weights4`, defined below.  Catch the errors generated in each case as an exception and print a useful message for the user.  

Hint: The first 3 test cases should return an exception!

Example:

```
weighted_avg(grades4, weights4) 

85.0
```

<!-- grade_id: cell-67380fa0fe8c0b2b, points: 0, grade: false, solution: true -->
```python
# Your code here
def weighted_avg(grades, weights):
    total_grade  = 0
    total_weight = 0
    try:    
        if len(grades) != len(weights):
            raise Exception("The number of weights and grades are not equal.")
        
        for k in range(0,len(grades)):
            if weights[k] < 0 or weights[k] >100:
                raise Exception("A weight is less than 0 or greater than 100.")
                
            if grades[k] < 0:
                raise Exception("A grade is below 0.")
                
            total_weight += weights[k]
            total_grade += grades[k]*weights[k]
        
        if total_weight != 100:
            raise Exception("Weights do not add up to 100.")
        
        return round(total_grade / 100,1)
    
    except Exception as e:
        return "An error has occurred: " + str(e)
    

grades1 = [88,99,100,70]
weights1 = [30, 30, 30, 5]

grades2 = [78, 75, 80, 99]
weights2 = [110, 10, -20, 0]

grades3 = [84, 80, 67, 97]
weights3 = [50, 25, 25]

grades4 = [100, 80, 90, 75]
weights4 = [20, 25, 25, 30] 
```

```python
# You can use the sample output in this cell to sanity check your work.

print(weighted_avg(grades4, weights4) == 85.0)
```

Output:
```
True
```

```python
# You can use the sample output in this cell to sanity check your work. 
# Does it throw an error? Is the error message useful?

print(weighted_avg(grades1, weights1))
```

Output:
```
An error has occurred: Weights do not add up to 100.
```

```python
# You can use the sample output in this cell to sanity check your work. 
# Does it throw an error? Is the error message useful?

print(weighted_avg(grades2, weights2))
```

Output:
```
An error has occurred: A weight is less than 0 or greater than 100.
```

```python
# You can use the sample output in this cell to sanity check your work. 
# Does it throw an error? Is the error message useful?

print(weighted_avg(grades3, weights3))
```

Output:
```
An error has occurred: The number of weights and grades are not equal.
```

<!-- grade_id: cell-429e12fb80ce822a, points: 5, grade: true, solution: false -->
```python
# Autograde cell - do not erase/delete
# Here is an example of how we autograde.

grades4 = [100, 80, 90, 75]
weights4 = [20, 25, 25, 30]


print(weighted_avg(grades4, weights4))
assert_equal(weighted_avg(grades4, weights4), 85.0)

```

<!-- grade_id: cell-95422c0a969fe204, points: 2, grade: true, solution: false -->
```python
# Hidden autograde cell - do not erase/delete
# Check your code for different grade and weight input variations as well! 
```

<!-- grade_id: cell-7dc17577706d51f5, points: 2, grade: true, solution: false -->
```python
# Hidden autograde cell - do not erase/delete
```

<!-- grade_id: cell-6f64ab4288079047, points: 1, grade: true, solution: false -->
```python
# Hidden autograde cell - do not erase/delete
```

<!-- grade_id: cell-4a23255307a51061, points: 1, grade: true, solution: false -->
```python
# Hidden autograde cell - do not erase/delete
```

<!-- grade_id: cell-70fe59b4d7ff1dcd, points: 2, grade: true, solution: false -->
```python
# Hidden autograde cell - do not erase/delete
```

## If you have feedback for this homework, please submit it using the link below:

http://goo.gl/forms/74yCiQTf6k