## Week 6 Assignment - W200 Python for Data Science, UC Berkeley MIDS

Write code in this Jupyter Notebook to solve each of the following problems. Each problem should have its solution in a separate cell. Please upload this **notebook**, your **scrabble.py** file, the **sowpods.txt** file, and your **score_word** module to your GitHub repository in your SUBMISSIONS/week_06 folder by 11:59PM PST the night before class.

## Objectives:

- Read and understand PEP 8 standards
- Use all of your previously gained knowledge together on a single program
- Demonstrate how to import a user made module and function into python from another .py file
- Demonstrate how to input command line arguments into a .py file

## 6-1. PEP 8 Style Guide (reading and response)

Your first task for this week is to write a **250 word reading response** to the article referenced below. In addition, please list **3 questions** that you have from the article. Please write your response in the markdown cell below.

The writing response is a free response, so you may write about your reactions, such as an interesting thing that you saw in the article, something that really stuck out to you, etc.

**Article**: [The PEP 8 Style Guide](https://www.python.org/dev/peps/pep-0008). This document is really important for Python coders because it describes best practices and customs.


<!-- grade_id: cell-eaae902b022d20da, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## 6-2. Cheating at Scrabble

Write a Python program that takes a Scrabble rack as a command-line argument and prints all valid Scrabble words that can be constructed from that rack, along with their Scrabble scores, sorted by score. Valid Scrabble words are provided in the data source below. A Scrabble rack is made up of 2 to 7 characters.

Below are the requirements for the program:
- This needs to be able to be run as a command line tool as shown below (not an input statement!)
- Name the python file: `scrabble.py`
- Include a function called `score_word` in a separate module. Import this function into your main solution code
- Allow anywhere from 2-7 character tiles (letters A-Z) to be inputted 
- Do not restrict the number of same tiles (e.g., a user is allowed to input ZZZZZQQ)
- Output the **total** list of valid Scrabble words that can be constructed from the rack as (score, word) tuples, sorted by score as shown below
- Then output 'Total number of words:' and the total number
- You need to handle input errors from the user and suggest what that error might be caused by and how to fix it (i.e., a helpful error message)
- Implement wildcards as either `*` or `?`. A wildcard character can take any value A-Z. There can be a total of **only** two wild cards in any user input (that is, one of each character: one `*` and one `?`). Only use the `*` and `?` as wildcard characters.
- Wildcard characters are scored as 0 points, just like in the real Scrabble game
- Your program should take less than two minutes to run with 2 wildcards in the input - if it takes more time than that then your algorithm is not optimized very well!
- Write docstrings for the functions and puts comments in your code.
- You should only use the Python standard library in this assignment.  However, any function in the standard library is allowed.

Extra Credit (+10 points):
Allow a user to specify that a certain letter has to be at a certain location. Your program must work without it so this is completely optional. For the extra credit, locations of certain letters must be specified at the command line, and may not be some sort of user prompt.  Please put comments and a sample of how to run your extra credit in the extra credit cell of this notebook below - it is the last cell.

An example invocation and output:
```
$ python scrabble.py ZAEFIEE
(17, feeze)
(17, feaze)
(16, faze)
(15, fiz)
(15, fez)
(12, zee)
(12, zea)
(11, za)
(6, fie)
(6, fee)
(6, fae)
(5, if)
(5, fe)
(5, fa)
(5, ef)
(2, ee)
(2, ea)
(2, ai)
(2, ae)
Total number of words: 19
```

The Data

http://courses.cms.caltech.edu/cs11/material/advjava/lab1/sowpods.zip contains all words in the official SOWPODS word list, one word per line. You should download the word file and keep it in your repository so that the program is standalone (instead of accessing it over the web from Python).

You can read data from a text file with the following code:

```
with open("sowpods.txt","r") as infile:
    raw_input = infile.readlines()
    data = [datum.strip('\n') for datum in raw_input]
```
This will show the first 6 words:
```
print(data[0:6])
```

Please use the dictionary below containing the letters and their Scrabble values:

```
scores = {"a": 1, "c": 3, "b": 3, "e": 1, "d": 2, "g": 2,
         "f": 4, "i": 1, "h": 4, "k": 5, "j": 8, "m": 3,
         "l": 1, "o": 1, "n": 1, "q": 10, "p": 3, "s": 1,
         "r": 1, "u": 1, "t": 1, "w": 4, "v": 4, "y": 4,
         "x": 8, "z": 10}
```

Tips:
- If you don't know what "scrabble" is or the basic background of the game please look it up online!
- We recommend that you try to break down the problem into steps on your own before writing any code. Once you've scoped generally what you want to do, then start writing some code.  If you get stuck, go back to thinking about the problem rather than trying to fix lots of errors at the code level.
- If you keep getting stuck, then check out: https://openhatch.org/wiki/Scrabble_challenge. This is where we got the idea for this assignment and it provides some helpful tips for guiding you along the way.  However, we would recommend that you try to implement this first before looking at the hints on the website.

Good luck!

### The code below will test your command line implementation of the scrabble.py code. We've made some of these tests available for you to try!

```python
# Code for the testing

import subprocess
from nose.tools import assert_equal 
from nose.tools import assert_true
from nose.tools import assert_greater
from nose.tools import assert_less
```

```python
""" Code runs and can produce at least one error message """

# no rack error
!python scrabble.py  
```

```python
""" Does not fail due to trivial mistakes and takes correct wildcard characters """

# does not fail due to case
!python scrabble.py "PENguin"

# takes wildcards
!python scrabble.py "PEN*?in"

```

```python
""" Produces a list of all words and scores that matches our expectations """

# test whether your output matches our expectation
cmd = ['python', 'scrabble.py', 'Penguin']
test = bytes.decode(subprocess.Popen(cmd, stdout=subprocess.PIPE).communicate()[0])
solution = '(10, penguin)\r\n(9, pening)\r\n(8, unpeg)\r\n(8, genip)\r\n(7, unpin)\r\n(7, unpen)\r\n(7, pung)\r\n(7, ping)\r\n(7, penni)\r\n(7, ingenu)\r\n(6, pug)\r\n(6, pine)\r\n(6, pig)\r\n(6, peni)\r\n(6, pein)\r\n(6, peg)\r\n(6, gup)\r\n(6, gip)\r\n(5, pun)\r\n(5, piu)\r\n(5, pin)\r\n(5, pie)\r\n(5, pen)\r\n(5, nip)\r\n(5, nep)\r\n(5, ginn)\r\n(5, gien)\r\n(5, genu)\r\n(5, ennui)\r\n(4, up)\r\n(4, pi)\r\n(4, pe)\r\n(4, nine)\r\n(4, neg)\r\n(4, gun)\r\n(4, gue)\r\n(4, gnu)\r\n(4, gin)\r\n(4, gie)\r\n(4, gen)\r\n(4, eng)\r\n(3, uni)\r\n(3, ug)\r\n(3, nun)\r\n(3, nie)\r\n(3, inn)\r\n(3, gu)\r\n(3, gi)\r\n(2, un)\r\n(2, nu)\r\n(2, ne)\r\n(2, in)\r\n(2, en)\r\nTotal number of words: 53\r\n'

print("\nDoes your output match our expectation?", test == solution)
```

<!-- grade_id: cell-ceab3d4433043319, points: 10, grade: true, solution: false -->
```python
""" The code should run in seconds to two minutes """

import time
start = time.time()

#test the code in the command line
cmd = ['python', 'scrabble.py', 'PENGU*?']
out = bytes.decode(subprocess.Popen(cmd, stdout=subprocess.PIPE).communicate()[0])

tot_time = time.time() - start
print('Total time was {} seconds'.format(tot_time))
assert_less(tot_time, 120)
```

<!-- grade_id: cell-7c027ce8309adb3d, points: 3, grade: true, solution: false -->
```python
# Autograding test
```

<!-- grade_id: cell-50add00e9627b5ac, points: 3, grade: true, solution: false -->
```python
# Autograding test
```

<!-- grade_id: cell-bab1478c7cd4a7f2, points: 3, grade: true, solution: false -->
```python
# Autograding test
```

<!-- grade_id: cell-eb305c64cbf12b26, points: 5, grade: true, solution: false -->
```python
# Autograding test
```

<!-- grade_id: cell-bb7a872eda80e6a5, points: 5, grade: true, solution: false -->
```python
# Autograding test
```

<!-- grade_id: cell-8d2de0735f148301, points: 5, grade: true, solution: false -->
```python
# Autograding test
```

<!-- grade_id: cell-16d4589caf032f51, points: 5, grade: true, solution: false -->
```python
# Autograding test
```

<!-- grade_id: cell-157fccf420d13535, points: 10, grade: true, solution: true -->
```python
"""Implement extra credit call run the code here
   If this cell isnt filled out - we'll assume the extra credit wasn't done
   Please write a comment on how to use your extra credit syntax also!
"""

# YOUR CODE HERE
```

## If you have feedback for this homework, please submit it using the link below:

http://goo.gl/forms/74yCiQTf6k